```python
%matplotlib inline
```

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from nose.tools import *

np.random.seed(24680)
```

Write your imports in the cell below.

<!-- grade_id: cell-98c4a3a245e45b8a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.preprocessing import MinMaxScaler

from sklearn.model_selection import train_test_split, StratifiedKFold, GridSearchCV

from sklearn.ensemble import RandomForestClassifier
from sklearn.tree import DecisionTreeClassifier
from sklearn.svm import LinearSVC, SVC

from sklearn.metrics import classification_report
#raise NotImplementedError()
```

# Ensemble Models and Support Vector Machines Lab
## Training and comparing different algorithms

Once again, we'll work with the bank dataset. This time, the data preprocessing steps have been done for you.

The goal is to try and improve our predictions (if they can be improved at all) using different types of algorithms.

### 1. Read the data (1 point)
This time you only need to read the data. The indicator variables have been separated out for you.

Read the dataset and save it in the variable `bank_data`. The target column is `y`. Use the variables `bank_attributes` and `bank_labels` to save the attributes (explanatory variables, features, predictors), and labels (`y`).

<!-- grade_id: cell-990a0dfca3a695e6, points: 0, grade: false, solution: true -->
```python
bank_data = pd.read_csv("data/bank.csv", delimiter = ",")
bank_attributes, bank_labels = bank_data.drop("y", axis = 1), bank_data["y"]
# YOUR CODE HERE
#raise NotImplementedError()
```

```python
bank_data
```

Output:
```
age  balance  day  duration  campaign  pdays  previous  job_admin.  \
0      30     1787   19        79         1     -1         0           0   
1      33     4789   11       220         1    339         4           0   
2      35     1350   16       185         1    330         1           0   
3      30     1476    3       199         4     -1         0           0   
4      59        0    5       226         1     -1         0           0   
...   ...      ...  ...       ...       ...    ...       ...         ...   
4516   33     -333   30       329         5     -1         0           0   
4517   57    -3313    9       153         1     -1         0           0   
4518   57      295   19       151        11     -1         0           0   
4519   28     1137    6       129         4    211         3           0   
4520   44     1136    3       345         2    249         7           0   

      job_blue-collar  job_entrepreneur  ...  month_mar  month_may  month_nov  \
0           
```

```python
bank_attributes.shape, bank_labels.shape
```

Output:
```
((4521, 51), (4521,))
```

<!-- grade_id: cell-dccbbb44b14b188f, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_data)
assert_is_not_none(bank_attributes)
assert_is_not_none(bank_labels)
```

### 2. Normalize the data (1 point)
Because both forests and SVMs are sensitive to non-scaled data, we need to normalize our dataset first.

Rescale all columns in `bank_attributes` so they have mean = 0 and variance = 1. You can either look at the `sklearn` docs or do this yourself. When you're ready, overwrite the `bank_attributes` column. Make sure that you don't lose the column names in the process.

<!-- grade_id: cell-7cc2e9c90f8705c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
scaler = MinMaxScaler()
bank_attributes = scaler.fit_transform(bank_attributes)
#raise NotImplementedError()
```

```python
bank_attributes
```

Output:
```
array([[0.16176471, 0.06845546, 0.6       , ..., 0.        , 0.        ,
        1.        ],
       [0.20588235, 0.10875022, 0.33333333, ..., 0.        , 0.        ,
        0.        ],
       [0.23529412, 0.06258976, 0.5       , ..., 0.        , 0.        ,
        0.        ],
       ...,
       [0.55882353, 0.04842888, 0.6       , ..., 0.        , 0.        ,
        1.        ],
       [0.13235294, 0.05973074, 0.16666667, ..., 1.        , 0.        ,
        0.        ],
       [0.36764706, 0.05971732, 0.06666667, ..., 1.        , 0.        ,
        0.        ]])
```

<!-- grade_id: cell-918a6111ee668331, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_attributes)
```

### 3. Split the data (1 point)
Use the standard 70% / 30% split. Since this is a classification problem, be sure to stratify the split according to the `bank_labels`.

<!-- grade_id: cell-68893dc2052dd87c, points: 0, grade: false, solution: true -->
```python
bank_attributes_train, bank_labels_train = None, None
bank_attributes_test, bank_labels_test = None, None
# YOUR CODE HERE
#raise NotImplementedError()
```

```python
bank_attributes_train, bank_attributes_test, bank_labels_train, bank_labels_test = train_test_split(
    bank_attributes,
    bank_labels,
    test_size=0.3,
    stratify=bank_labels
)
```

<!-- grade_id: cell-215de695ca1c45ff, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_attributes_train)
assert_is_not_none(bank_labels_train)

assert_is_not_none(bank_attributes_test)
assert_is_not_none(bank_labels_test)
```

### 4. Prepare the cross-validation folds (1 point)
Use a stratified k-fold cross-validation split, with $k = 5$. Fit it to the train data. Save the trained cross-validator to the variable `k_fold`.

The data should already be shuffled. There's no need to shuffle it again.

<!-- grade_id: cell-ef1a3912d43342f1, points: 0, grade: false, solution: true -->
```python
k_fold = None
# YOUR CODE HERE
skf = StratifiedKFold(n_splits=5)
#raise NotImplementedError()
```

```python
skf.get_n_splits(bank_attributes_train, bank_labels_train)
```

Output:
```
5
```

```python
for train_index, test_index in skf.split(bank_attributes_train, bank_labels_train):
    print("Train Index Length: ", len(train_index), "Test Index Length: ", len(test_index))
    #print("TRAIN:", train_index, "Having: ", len(train_index), "TEST:", test_index, "Having: ", len(test_index))
    #X_train, X_test = bank_attributes_train[train_index], bank_attributes_train[test_index]
    #y_train, y_test = bank_labels_train[train_index], bank_labels_train[test_index]
```

Output:
```
Train Index Length:  2531 Test Index Length:  633
Train Index Length:  2531 Test Index Length:  633
Train Index Length:  2531 Test Index Length:  633
Train Index Length:  2531 Test Index Length:  633
Train Index Length:  2532 Test Index Length:  632
```

```python
#k_fold = skf.split(bank_attributes_train, bank_labels_train)
```

```python
k_fold = StratifiedKFold(n_splits=5)
```

<!-- grade_id: cell-8c28c6c421c7705a, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(k_fold)
```

### 5. Decision Tree (2 points)
Use cross-validation to train and optimize the hyperparameters for a decision tree classifier.

Use grid search with the following grid:
* `max_depth`: 1, 5, 7, 15, 20
* `min_samples_leaf`: 2, 5, 10, 12
* `max_leaf_nodes`: 5, 10, 20

Use the most appropriate scoring metric (remember that accuracy doesn't work in this case because the data is highly imbalanced; we need something which combines precision and recall). Use the cross-validation splits you just created.

Save the grid results in `grid_search`. Save the best classifier in `tree_classifier`.

Optionally, you can print and / or visualize the cross-validation results and the best chosen parameters.

```python
params = {
    "max_depth": [20, 50, 100],
    "min_samples_leaf": [2, 5, 10, 12],
    "max_leaf_nodes": [5, 10, 20]
}

grid_search = GridSearchCV(estimator = DecisionTreeClassifier(), param_grid=params, scoring='f1', cv=k_fold)
#for train_index, test_index in k_fold:
grid_search.fit(bank_attributes_train, bank_labels_train)
```

Output:
```
GridSearchCV(cv=StratifiedKFold(n_splits=5, random_state=None, shuffle=False),
             estimator=DecisionTreeClassifier(),
             param_grid={'max_depth': [20, 50, 100],
                         'max_leaf_nodes': [5, 10, 20],
                         'min_samples_leaf': [2, 5, 10, 12]},
             scoring='f1')
```

```python
grid_search.best_estimator_
```

Output:
```
DecisionTreeClassifier(max_depth=20, max_leaf_nodes=10, min_samples_leaf=2)
```

```python
#grid_search.cv_results_
```

```python
tree = DecisionTreeClassifier(max_depth=20, max_leaf_nodes=10, min_samples_leaf=2)
tree.fit(bank_attributes_train, bank_labels_train)
```

Output:
```
DecisionTreeClassifier(max_depth=20, max_leaf_nodes=10, min_samples_leaf=2)
```

<!-- grade_id: cell-3603341dc3a0e1b4, points: 0, grade: false, solution: true -->
```python
grid_search = grid_search
tree_classifier = tree

# YOUR CODE HERE

#raise NotImplementedError()

print("Decision tree; best score:", grid_search.best_score_)
```

Output:
```
Decision tree; best score: 0.5090152112059995
```

<!-- grade_id: cell-8a9d43f3f99d1c09, points: 2, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(tree_classifier)
```

### 5. Random Forest (1 point)
Use cross-validation to train and optimize the hyperparameters for a random forest classifier. Use the same technique as before.

Use the following grid:
* `n_estimators`: 100, 200, 300 
* `max_depth`: 20, 50, 100

Note that this grid is on the small side but this is mainly due to performance reasons. Also note that the training will take some time.

Save the grid results in `grid_search`. Save the best classifier in `forest_classifier`.

Optionally, you can print and / or visualize the cross-validation results and the best chosen parameters.

Due to the relatively slow training, we've chosen low values for the parameters. The performance of the random forest will be worse than the decision tree. This is not necessarily the case in general, it's due to the parameters we've chosen to try.

```python
params = {
    "n_estimators": [100, 200, 300],
    "max_depth": [20, 50, 100]
}
grid_search = GridSearchCV(estimator = RandomForestClassifier(), param_grid=params, scoring='f1', cv=k_fold)
grid_search.fit(bank_attributes_train, bank_labels_train)
```

Output:
```
GridSearchCV(cv=StratifiedKFold(n_splits=5, random_state=None, shuffle=False),
             estimator=RandomForestClassifier(),
             param_grid={'max_depth': [20, 50, 100],
                         'n_estimators': [100, 200, 300]},
             scoring='f1')
```

```python
grid_search.best_estimator_
```

Output:
```
RandomForestClassifier(max_depth=50, n_estimators=300)
```

```python
#grid_search.cv_results_
```

```python
run_forest_run = RandomForestClassifier(max_depth=100, n_estimators=200)
run_forest_run.fit(bank_attributes_train, bank_labels_train)
```

Output:
```
RandomForestClassifier(max_depth=100, n_estimators=200)
```

```python
#run_forest_run.fit(bank_attributes_train, bank_labels_train)
```

<!-- grade_id: cell-965717809a95920b, points: 0, grade: false, solution: true -->
```python
grid_search = grid_search
forest_classifier = run_forest_run


# YOUR CODE HERE
#raise NotImplementedError()

print("Random forest; best score:", grid_search.best_score_)
```

Output:
```
Random forest; best score: 0.3884528136817182
```

<!-- grade_id: cell-5cdc884e24e42fec, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(forest_classifier)
```

### 6. Linear SVM (1 point)
Use cross-validation to train and optimize the hyperparameters for a linear support vector machine. Use the same technique as before.

Use the following grid:
* `C`: 0.1, 0.5, 0.8, 1, 1.5, 2, 6, 10, 15, 20

Note that we're choosing relatively small values for `C`. This is allowed because our data is normalized.

Save the grid results in `grid_search`. Save the best classifier in `linear_svm_classifier`. There are many ways to create a linear SVM classifier. Look up the `sklearn` docs to choose the fastest one (in terms of performance).

Optionally, you can print and / or visualize the cross-validation results and the best chosen parameters.

```python
params = {
    "C": [0.1, 0.5, 0.8, 1, 1.5, 2, 6, 10, 15, 20]
}
grid_search = GridSearchCV(estimator = LinearSVC(), param_grid=params, scoring='f1', cv=k_fold)
grid_search.fit(bank_attributes_train, bank_labels_train)
```

Output:
```
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\svm\_base.py:985: ConvergenceWarning: Liblinear failed to converge, increase the number of iterations.
  warnings.warn("Liblinear failed to converge, increase "
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\svm\_base.py:985: ConvergenceWarning: Liblinear failed to converge, increase the number of iterations.
  warnings.warn("Liblinear failed to converge, increase "
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\svm\_base.py:985: ConvergenceWarning: Liblinear failed to converge, increase the number of iterations.
  warnings.warn("Liblinear failed to converge, increase "
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\svm\_base.py:985: ConvergenceWarning: Liblinear failed to converge, increase the number of iterations.
  warnings.warn("Liblinear failed to converge, increase "
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\svm\_base.py:985: ConvergenceWarning: Liblinear failed to converge, increase the number of iterations
```

```python
grid_search.best_params_
```

Output:
```
{'C': 20}
```

```python
linearSVC = LinearSVC(C=20)
linearSVC.fit(bank_attributes_train, bank_labels_train)
```

Output:
```
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\svm\_base.py:985: ConvergenceWarning: Liblinear failed to converge, increase the number of iterations.
  warnings.warn("Liblinear failed to converge, increase "
LinearSVC(C=20)
```

```python
#linearSVC.fit(bank_attributes_train, bank_labels_train)
```

<!-- grade_id: cell-dbc74911ea58e898, points: 0, grade: false, solution: true -->
```python
grid_search = grid_search
linear_svm_classifier = linearSVC

# YOUR CODE HERE
#raise NotImplementedError()

print("Linear SVM; best score:", grid_search.best_score_)
```

Output:
```
Linear SVM; best score: 0.39103350589477115
```

<!-- grade_id: cell-7e5b16fe3d5deb03, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(linear_svm_classifier)
```

### 7. Gaussian SVM (1 point)
Use cross-validation to train and optimize the hyperparameters for an SVM with a Gaussian kernel. Use the same technique as before.

Use the following grid:
* `C`: 10, 15, 20, 50, 200
* `gamma`: 0.001, 0.01, 0.1, 0.2

Note that this time we give larger values of `C` because the governing parameter here is `gamma`.

Save the grid results in `grid_search`. Save the best classifier in `gaussian_svm_classifier`.

Optionally, you can print and / or visualize the cross-validation results and the best chosen parameters.

```python
params = {
    "C": [10, 15, 20, 50, 200],
    "C": [0.001, 0.01, 0.1, 0.2]
}
grid_search = GridSearchCV(estimator = SVC(kernel="rbf"), param_grid=params, scoring='f1', cv=k_fold)
grid_search.fit(bank_attributes_train, bank_labels_train)
```

Output:
```
GridSearchCV(cv=StratifiedKFold(n_splits=5, random_state=None, shuffle=False),
             estimator=SVC(), param_grid={'C': [0.001, 0.01, 0.1, 0.2]},
             scoring='f1')
```

```python
grid_search.best_estimator_
```

Output:
```
SVC(C=0.001)
```

```python
#grid_search.cv_results_
```

```python
gaussian_svm = SVC(kernel = "rbf", C = 0.001)
gaussian_svm.fit(bank_attributes_train, bank_labels_train)
```

Output:
```
SVC(C=0.001)
```

<!-- grade_id: cell-06ab026bf98b2e69, points: 0, grade: false, solution: true -->
```python
grid_search = grid_search
gaussian_svm_classifier = gaussian_svm

# YOUR CODE HERE
#raise NotImplementedError()

print("Gaussian SVM; best score:", grid_search.best_score_)
```

Output:
```
Gaussian SVM; best score: 0.0
```

<!-- grade_id: cell-0ba2a8ae7d45a6ad, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(gaussian_svm_classifier)
```

### 9. Compare performance on the testing data (1 point)
Now that you've trained all your models, you've got to select the best one. This should be done on the testing data.

Use the appropriate scoring metric to get the testing scores for all your models. Don't forget to pass the **testing**, not the training data. Save all scores.

Choose the best classifier, based on these scores (the one with the highest test score). Of course, this is not enough. We need to look at ROC curves, track performance through other measures, debug the sources of variance in testing results, try more hyperparameters, etc. However, this is enough for an introductory lab :).

Optionally, you can think of combining them into a boosted model but this is out of the scope of this lab.

```python
tree_classifier.score(bank_attributes_test, bank_labels_test)
```

Output:
```
0.8953574060427414
```

```python
print(classification_report(bank_labels_test,
                              forest_classifier.predict(bank_attributes_test),
                              output_dict=True)['macro avg']['f1-score'])
```

Output:
```
0.6488044032951812
```

```python
algs = [tree_classifier, forest_classifier, linear_svm_classifier, gaussian_svm_classifier]
scores = []
for alg in algs:
    scores.append(
        classification_report(bank_labels_test,
                              alg.predict(bank_attributes_test),
                              output_dict=True)['macro avg']['f1-score']
    )
```

Output:
```
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\metrics\_classification.py:1245: UndefinedMetricWarning: Precision and F-score are ill-defined and being set to 0.0 in labels with no predicted samples. Use `zero_division` parameter to control this behavior.
  _warn_prf(average, modifier, msg_start, len(result))
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\metrics\_classification.py:1245: UndefinedMetricWarning: Precision and F-score are ill-defined and being set to 0.0 in labels with no predicted samples. Use `zero_division` parameter to control this behavior.
  _warn_prf(average, modifier, msg_start, len(result))
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\metrics\_classification.py:1245: UndefinedMetricWarning: Precision and F-score are ill-defined and being set to 0.0 in labels with no predicted samples. Use `zero_division` parameter to control this behavior.
  _warn_prf(average, modifier, msg_start, len(result))
```

```python
print(scores)
```

Output:
```
[0.7020810142238714, 0.6488044032951812, 0.6747860457361972, 0.46950742767787335]
```

<!-- grade_id: cell-cf416134ae2b2a62, points: 0, grade: false, solution: true -->
```python
tree_classifier_score, forest_classifier_score, linear_svm_classifier_score, gaussian_svm_classifier_score = scores
# YOUR CODE HERE
#raise NotImplementedError()

print("Testing scores:")
print("Decision tree:", tree_classifier_score)
print("Random forest:", forest_classifier_score)
print("Linear SVM:", linear_svm_classifier_score)
print("Gaussian SVM:", gaussian_svm_classifier_score)
```

Output:
```
Testing scores:
Decision tree: 0.7020810142238714
Random forest: 0.6488044032951812
Linear SVM: 0.6747860457361972
Gaussian SVM: 0.46950742767787335
```

<!-- grade_id: cell-bdeed4c0d761fe0e, points: 0, grade: false, solution: true -->
```python
best_classifier = "tree" # Replace empty string with "tree", "forest", "linear SVM" or "gaussian SVM"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-495fc5ef0619deb6, points: 1, grade: true, solution: false -->
```python
assert_not_equal(best_classifier, "")
```

```python

```