```python
%matplotlib inline
```

```python
from nose.tools import *

# Write your imports in the cell below
```

<!-- grade_id: cell-3a5c071a139ade21, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.preprocessing import MinMaxScaler

from sklearn.model_selection import train_test_split, GridSearchCV

from sklearn.linear_model import LogisticRegression

from sklearn.metrics import classification_report

#raise NotImplementedError()
```

```python
np.random.seed(1234)
```

# Model Training and Improvement Lab
## Comparing and selecting models

### 1. Read the data (1 point)
Like in the previous lab, you need to read the Portuguese bank dataset [here](https://archive.ics.uci.edu/ml/machine-learning-databases/00222/). It has been provided for you in the `data` folder.

Read the dataset using `pandas` (you can use the library with the alias `pd`). Save it in the `bank_data` variable.

<!-- grade_id: cell-39fe3843121d7b36, points: 0, grade: false, solution: true -->
```python
bank_data = pd.read_csv("data/bank.csv", delimiter = ";")
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-b183f83b98f16a29, points: 1, grade: true, solution: false -->
```python
# From now on, all test cells might contain hidden tests. If you follow the instructions correctly, 
# your solution will be graded with maximum points
assert_is_not_none(bank_data)
```

### 2. Preprocess the data (1 point)
Separate explanatory features from labels. Save all features (16 columns total) in the variable `bank_features`. Save the labels (corresponding to the `y` column) in the `bank_labels` variable. Rewrite the labels to be `0` and `1` instead of `no` and `yes`: `bank_labels` should be a numeric column.

<!-- grade_id: cell-fb1c7a38947afc06, points: 0, grade: false, solution: true -->
```python
bank_features, bank_labels = bank_data.drop("y", axis = 1), bank_data["y"]
# YOUR CODE HERE
#raise NotImplementedError()
```

```python
bank_features.shape, bank_labels.shape
```

Output:
```
((4521, 16), (4521,))
```

<!-- grade_id: cell-30a215c478663fb9, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features)
assert_is_not_none(bank_labels)
```

### 3. Get indicator variables (1 point)
Get indicator (dummy) variables for all categorical columns in `bank_features`. Overwrite the `bank_features` variable to store the new data.

<!-- grade_id: cell-8724b6857c2eb7fc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
bank_features.dtypes
#raise NotImplementedError()
```

Output:
```
age           int64
job          object
marital      object
education    object
default      object
balance       int64
housing      object
loan         object
contact      object
day           int64
month        object
duration      int64
campaign      int64
pdays         int64
previous      int64
poutcome     object
dtype: object
```

```python
bank_features = pd.get_dummies(bank_features)
```

```python
bank_features.shape
```

Output:
```
(4521, 51)
```

<!-- grade_id: cell-c2b2c0bd2d30ea5e, points: 1, grade: true, solution: false -->
```python
assert_equal(bank_features.shape, (4521, 51))
```

### 4. Split the data (1 point)
Split the data into training and testing set, with 70% of the data for training. Because the output labels are not equaly distributed, use stratification based on the `bank_labels`.

<!-- grade_id: cell-f46c32bb67654092, points: 0, grade: false, solution: true -->
```python
bank_features_train, bank_labels_train = None, None
bank_features_test, bank_labels_test = None, None
# YOUR CODE HERE
bank_features_train, bank_features_test, bank_labels_train, bank_labels_test = train_test_split(
bank_features, bank_labels, train_size=0.7, random_state=42, stratify=bank_labels)
#raise NotImplementedError()
```

<!-- grade_id: cell-b14f41fb226c0f37, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features_train)
assert_is_not_none(bank_labels_train)
assert_is_not_none(bank_features_test)
assert_is_not_none(bank_labels_test)
```

### 5. Train a baseline algorithm (1 point)
Train a logistic regression using the training data. Use 1 000 000 (`1e6`) as the value of C. Score it using the testing data. Save the score in the `baseline_score` variable. You should see a fairly high score.

<!-- grade_id: cell-2fd7633e6266fcba, points: 0, grade: false, solution: true -->
```python
model = None
baseline_score = None

# YOUR CODE HERE
#raise NotImplementedError()
```

```python
model = LogisticRegression(C=1e6)
model.fit(bank_features_train, bank_labels_train)
```

Output:
```
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\linear_model\_logistic.py:763: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
LogisticRegression(C=1000000.0)
```

```python
baseline_score = model.score(bank_features_test, bank_labels_test)
print(baseline_score)
```

Output:
```
0.89240972733972
```

<!-- grade_id: cell-a3af8c95790cccc4, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(model)
assert_greater(baseline_score, 0.7)
```

### 6. Select a better score (2 points)
As you alrady saw, the positive examples are very few. If you aren't convinced, just check the counts.

We know that the default scoring (accuracy) isn't correct in this case. Better measures would be precision and recall. However, we only want one number. Evaluate the algorithm once again, using a standard scoring method which combines precision and recall. Overwrite the `baseline_score` variable.

Don't forget to score the model on the testing data only.

<!-- grade_id: cell-9b44f24da39ba641, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
bank_predictions_test = model.predict(bank_features_test)
baseline_score = classification_report(bank_labels_test, bank_predictions_test, output_dict=True)
#raise NotImplementedError()
baseline_score = baseline_score['macro avg']['f1-score']
print(baseline_score)
```

Output:
```
0.6264837186854291
```

<!-- grade_id: cell-9f4aca08f036d9de, points: 2, grade: true, solution: false -->
```python
assert_less(baseline_score, 0.7)
```

### 7. Tune your model (2 points)
Fine-tune the `C` and `max_iter` parameters.

Use full grid search with the following values:
* `C`: 0.0001, 0.01, 0.1, 1, 10, 100, 10000
* `max_iter`: 50, 100, 300, 1000
* `fit_itercept`: True, False

Save the grid search result in the `grid_search` variable. Don't forget to use the better scoring model that you obtained in the previous task.

<!-- grade_id: cell-88cb1341dfd72476, points: 0, grade: false, solution: true -->
```python
params = {
    "C": [0.0001, 0.01, 0.1, 1, 10, 100, 10000],
    "max_iter": [50, 100, 300, 1000],
    "fit_intercept": [True, False]
}

grid_search = GridSearchCV(estimator = LogisticRegression(), param_grid=params, scoring='f1')
# YOUR CODE HERE
grid_search.fit(bank_features_train, bank_labels_train)
#raise NotImplementedError()
```

Output:
```
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\linear_model\_logistic.py:763: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\model_selection\_validation.py:683: UserWarning: Scoring failed. The score on this train-test partition for these parameters will be set to nan. Details: 
Traceback (most recent call last):
  File "C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\model_selection\_validation.py", line 674, in _score
    scores = scorer(estimator, X_test, y_test)
  File "C:\Users\ksbon\anaconda3\lib\site-packages\sklearn\metrics\_scorer.py", line 1
```

```python
grid_search.best_estimator_
```

Output:
```
LogisticRegression(C=0.0001, max_iter=50)
```

```python
grid_search.best_estimator_.score(bank_features_test, bank_labels_test)
```

Output:
```
0.8843036109064112
```

```python
grid_search.cv_results_
```

Output:
```
{'mean_fit_time': array([0.01941767, 0.03242912, 0.07086601, 0.07126241, 0.0192173 ,
        0.023421  , 0.02362132, 0.02362156, 0.01841645, 0.03202896,
        0.08687878, 0.27915959, 0.01821623, 0.03162866, 0.07767034,
        0.09408751, 0.01881723, 0.03203087, 0.09148283, 0.27925339,
        0.01781592, 0.03143549, 0.08487673, 0.26664195, 0.01861682,
        0.03182898, 0.08667884, 0.27384844, 0.01861677, 0.03143039,
        0.08567767, 0.27424707, 0.0188168 , 0.03222909, 0.08748112,
        0.27565012, 0.0184165 , 0.03163066, 0.08667679, 0.27345014,
        0.01861682, 0.03162837, 0.08567762, 0.27565007, 0.01821637,
        0.03162885, 0.08498192, 0.27064548, 0.01881685, 0.03260326,
        0.0860795 , 0.27765007, 0.01861677, 0.03162856, 0.08508081,
        0.27084579]),
 'std_fit_time': array([0.00049054, 0.00135777, 0.00768466, 0.00794085, 0.0007487 ,
        0.0024187 , 0.00162625, 0.00215602, 0.00049004, 0.00089527,
        0.00171904, 0.00201777, 0.00074913, 0.00049041, 0.015
```

<!-- grade_id: cell-0e7311dbe13cf894, points: 2, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(grid_search.best_estimator_)
```

### 8. Compare scores (1 point)
Use the best estimator from your grid search. Score it using the function from problem 6. Save your answer in `tuned_score`.

<!-- grade_id: cell-98568f584f6c0778, points: 0, grade: false, solution: true -->
```python
tuned_score = classification_report(
    bank_labels_test,
    grid_search.best_estimator_.predict(bank_features_test), output_dict=True)['macro avg']['f1-score']
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-545b4d41fef4eeef, points: 1, grade: true, solution: false -->
```python
print(tuned_score)
```

Output:
```
0.5621258115765395
```

```python
print(baseline_score - tuned_score)
```

Output:
```
0.06435790710888956
```

Hmmmm, it seems we have not obtained a better algorithm, even the opposite (the difference is marginal and depends on the random initialization of the cross-validation datasets).

We can, of course, do a lot more things to improve our model's performance, such as normalizing the data, feature selection and feature engineering, trying out different aspects, e.g. polynomial terms, RANSAC; even boosting (we'll talk about this later). However, we'll stop at this point.

What can we conclude? It seems that this is close to the best performance we can get out of this algorithm, given these data points.

We can try improving (cleaning) our dataset, selecting features, etc. but we most likely need a better algorithm. In the next labs, we're going to explore that.